<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Electrochemistry module tests</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<style>
.nuc-wrap {
	max-width: 1400px;
	font-size: 14px;
	line-height: 1.4;
}

.nuc-table {
	width: 100%;
	border-collapse: collapse;
	table-layout: fixed;
	margin: 12px 0 24px;
	background: #ffffff;
}

.nuc-table th,
.nuc-table td {
	border: 1px solid #cfd8dc;
	padding: 8px 10px;
	vertical-align: top;
	text-align: left;
}

.nuc-table th {
	background: #eef3f6;
	font-weight: 600;
}

.nuc-table tr:nth-child(even) td {
	background: #fafcfd;
}

.nuc-id {
	width: 70px;
}

.nuc-desc {
	width: 30%;
}

.nuc-result {
	width: 90px;
	font-weight: 600;
}

.nuc-output {
	width: 50%;
	white-space: pre-wrap;
	overflow-wrap: anywhere;
	word-break: break-word;
	font-family: Consolas, "Liberation Mono", monospace;
	font-size: 13px;
}

.nuc-summary {
	width: 100%;
	max-width: 520px;
	border-collapse: collapse;
	margin-top: 8px;
	background: #ffffff;
}

.nuc-summary td {
	border: 1px solid #cfd8dc;
	padding: 10px 12px;
}
</style>

<div class="nuc-wrap">
<h2>Electrochemistry Module Tests</h2>

<table class="nuc-table">
	<tr style="background-color:#f0f0f0;">
		<th class="nuc-id">Test ID</th>
		<th class="nuc-desc">Description</th>
		<th class="nuc-result">Result</th>
		<th class="nuc-output">Output</th>
	</tr>
	<tr><td>1</td><td><code>chem_electrochem_E0("Fe^{3+}","Fe^{2+}")</code></td><td>{@p_1@}</td><td class="nuc-output">{@t_1@}</td></tr>
	<tr><td>2</td><td><code>chem_electrochem_n("Fe^{3+}","Fe^{2+}")</code></td><td>{@p_2@}</td><td class="nuc-output">{@t_2@}</td></tr>
	<tr><td>3</td><td><code>chem_electrochem_E0("Zn^{2+}","Zn")</code></td><td>{@p_3@}</td><td class="nuc-output">{@t_3@}</td></tr>
	<tr><td>4</td><td><code>chem_electrochem_reductant("Fe^{3+}")</code></td><td>{@p_4@}</td><td class="nuc-output">{@t_4@}</td></tr>
	<tr><td>5</td><td><code>chem_electrochem_oxidant("Fe^{2+}")</code></td><td>{@p_5@}</td><td class="nuc-output">{@t_5@}</td></tr>
	<tr><td>6</td><td><code>chem_electrochem_pair_data("Fe^{3+}","Fe^{2+}")</code></td><td>{@p_6@}</td><td class="nuc-output">{@t_6@}</td></tr>
	<tr><td>7</td><td><code>chem_electrochem_half_equation("Fe^{3+}","Fe^{2+}")</code></td><td>{@p_7@}</td><td class="nuc-output">{@t_7@}</td></tr>
	<tr><td>8</td><td><code>chem_electrochem_oxidant_array()</code></td><td>{@p_8@}</td><td class="nuc-output">{@length(t_8)@}</td></tr>
	<tr><td>9</td><td><code>chem_electrochem_reductant_array()</code></td><td>{@p_9@}</td><td class="nuc-output">{@length(t_9)@}</td></tr>
	<tr><td>10</td><td><code>chem_electrochem_pairs_array()</code></td><td>{@p_10@}</td><td class="nuc-output">{@length(t_10)@}</td></tr>
	<tr><td>11</td><td><code>chem_electrochem_strong_oxidant_array(1.0)</code></td><td>{@p_11@}</td><td class="nuc-output">{@t_11@}</td></tr>
	<tr><td>12</td><td><code>chem_electrochem_strong_reductant_array(-2.0)</code></td><td>{@p_12@}</td><td class="nuc-output">{@t_12@}</td></tr>
	<tr><td>13</td><td><code>chem_electrochem_cell_potential(...)</code></td><td>{@p_13@}</td><td class="nuc-output">{@t_13@}</td></tr>
	<tr><td>14</td><td><code>chem_electrochem_spontaneous(...)</code></td><td>{@p_14@}</td><td class="nuc-output">{@t_14@}</td></tr>
	<tr><td>15</td><td><code>chem_electrochem_E0("X","Y")</code></td><td>{@p_15@}</td><td class="nuc-output">{@t_15@}</td></tr>
	<tr><td>16</td><td><code>chem_electrochem_array()</code></td><td>{@p_16@}</td><td class="nuc-output">{@length(t_16)@}</td></tr>
	<tr><td>17</td><td><code>chem_electrochem_half_equation_excl_H_H2O(...)</code></td><td>{@p_17@}</td><td class="nuc-output">{@t_17@}</td></tr>
	<tr><td>18</td><td><code>chem_electrochem_core_pair_data(...)</code></td><td>{@p_18@}</td><td class="nuc-output">{@t_18@}</td></tr>
	<tr><td>19</td><td><code>chem_electrochem_Hplus_species(...)</code></td><td>{@p_19@}</td><td class="nuc-output">{@t_19@}</td></tr>
	<tr><td>20</td><td><code>chem_electrochem_Hplus_count(...)</code></td><td>{@p_20@}</td><td class="nuc-output">{@t_20@}</td></tr>
	<tr><td>21</td><td><code>chem_electrochem_Hplus_count("Fe^{3+}","Fe^{2+}")</code></td><td>{@p_21@}</td><td class="nuc-output">{@t_21@}</td></tr>
</table>

<table class="nuc-summary">
	<tr><td>Total Tests</td><td>{@max_score@}</td></tr>
	<tr><td>Tests Passed</td><td>{@total_score@}</td></tr>
	<tr><td>Overall Score</td><td>{@score_pct@}%</td></tr>
</table>

<p>[[input:ans1]][[validation:ans1]]</p>
[[feedback:prt1]]
</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* electrochemistry tests */
        /* Self-contained electrochemistry module.
         Do not use stack_include here: Moodle's CAS sandbox can time out while
         fetching a remote GitHub URL. */
        %_ELECTROCHEM_DATA: [
          "F2", "2F-", 2, 2.87,
          "Ce^{4+}", "Ce^{3+}", 1, 1.61,
          "MnO4^- + 8H+", "Mn^{2+} + 4H2O", 5, 1.51,
          "Cl2", "2Cl-", 2, 1.36,
          "Cr2O7^{2-} + 14H+", "2Cr^{3+} + 7H2O", 6, 1.33,
          "O2 + 4H+", "2H2O", 4, 1.23,
          "MnO2 + 4H+", "Mn^{2+} + 2H2O", 2, 1.23,
          "Br2", "2Br-", 2, 1.06,
          "NO3^- + 4H+", "NO + 2H2O", 3, 0.96,
          "Ag+", "Ag", 1, 0.80,
          "Fe^{3+}", "Fe^{2+}", 1, 0.77,
          "O2 + 2H+", "H2O2", 2, 0.68,
          "MnO4^- + 2H2O", "MnO2 + 4OH-", 3, 0.59,
          "I2", "2I-", 2, 0.54,
          "O2 + 2H2O", "4OH-", 4, 0.40,
          "[Fe(CN)6]^{3-}", "[Fe(CN)6]^{4-}", 1, 0.370,
          "Cu^{2+}", "Cu", 2, 0.34,
          "AgCl", "Ag + Cl-", 1, 0.222,
          "H+", "0.5H2", 2, 0.00,
          "Pb^{2+}", "Pb", 2, -0.13,
          "Sn^{2+}", "Sn", 2, -0.14,
          "Ni^{2+}", "Ni", 2, -0.28,
          "Cd^{2+}", "Cd", 2, -0.40,
          "Fe^{2+}", "Fe", 2, -0.44,
          "Zn^{2+}", "Zn", 2, -0.76,
          "2H2O", "H2 + 2OH-", 2, -0.83,
          "Al^{3+}", "Al", 3, -1.66,
          "Mg^{2+}", "Mg", 2, -2.37,
          "Na+", "Na", 1, -2.71,
          "Ca^{2+}", "Ca", 2, -2.87,
          "Ba^{2+}", "Ba", 2, -2.90,
          "K+", "K", 1, -2.93,
          "Li+", "Li", 1, -3.05,
          "Co^{2+}", "Co", 2, -0.28,
          "Mn^{2+}", "Mn", 2, -1.18,
          "Cu+", "Cu", 1, 0.52,
          "H2O2 + 2H+", "2H2O", 2, 1.77,
          "PbO2 + 4H+", "Pb^{2+} + 2H2O", 1, 1.46,
          "Au^{3+}", "Au", 3, 1.41,
          "ClO4^- + 2H+", "ClO3^- + H2O", 2, 1.00,
          "NO3^- + 3H+", "HNO2 + H2O", 2, 0.94,
          "Hg^{2+}", "Hg", 2, 0.85,
          "Hg2Cl2", "2Hg + 2Cl-", 2, 0.27,
          "SO4^{2-} + 4H+", "H2SO3 + H2O", 2, 0.20,
          "Sn^{4+}", "Sn^{2+}", 2, 0.15,
          "HgO + H2O", "Hg + 2OH-", 2, 0.10,
          "Fe^{3+}", "Fe", 3, -0.036,
          "Zn(OH)2", "Zn + 2OH-", 2, -1.25
        ]$

        _electrochem_normalize_species(s) := block([x],
          x: ssubst("H+", "H^+", s),
          ssubst("OH-", "OH^-", x)
        )$

        _electrochem_contains(haystack, needle) := block([p],
          p: errcatch(ssearch(needle, haystack)),
          if is(length(p) = 0) then false else is(first(p) # false)
        )$

        _electrochem_find_oxidant(oxidant) := block([i, n, result],
          oxidant: _electrochem_normalize_species(oxidant),
          n: length(%_ELECTROCHEM_DATA), i: 1, result: 0,
          while i <= n and is(result = 0) do (
            if sequal(%_ELECTROCHEM_DATA[i], oxidant) then result: i,
            i: i + 4
          ),
          result
        )$

        _electrochem_find_reductant(reductant) := block([i, n, result],
          reductant: _electrochem_normalize_species(reductant),
          n: length(%_ELECTROCHEM_DATA), i: 1, result: 0,
          while i <= n and is(result = 0) do (
            if sequal(%_ELECTROCHEM_DATA[i + 1], reductant) then result: i,
            i: i + 4
          ),
          result
        )$

        _electrochem_find_pair(oxidant, reductant) := block([i, n, result],
          oxidant: _electrochem_normalize_species(oxidant),
          reductant: _electrochem_normalize_species(reductant),
          n: length(%_ELECTROCHEM_DATA), i: 1, result: 0,
          while i <= n and is(result = 0) do (
            if sequal(%_ELECTROCHEM_DATA[i], oxidant) and sequal(%_ELECTROCHEM_DATA[i + 1], reductant) then result: i,
            i: i + 4
          ),
          result
        )$

        chem_electrochem_E0(oxidant, reductant) := block([idx],
          idx: _electrochem_find_pair(oxidant, reductant),
          if is(idx = 0) then null else %_ELECTROCHEM_DATA[idx + 3]
        )$

        chem_electrochem_n(oxidant, reductant) := block([idx],
          idx: _electrochem_find_pair(oxidant, reductant),
          if is(idx = 0) then null else %_ELECTROCHEM_DATA[idx + 2]
        )$

        chem_electrochem_reductant(oxidant) := block([idx],
          idx: _electrochem_find_oxidant(oxidant),
          if is(idx = 0) then "" else %_ELECTROCHEM_DATA[idx + 1]
        )$

        chem_electrochem_oxidant(reductant) := block([idx],
          idx: _electrochem_find_reductant(reductant),
          if is(idx = 0) then "" else %_ELECTROCHEM_DATA[idx]
        )$

        chem_electrochem_pair_data(oxidant, reductant) := block([idx],
          idx: _electrochem_find_pair(oxidant, reductant),
          if is(idx = 0) then []
          else [%_ELECTROCHEM_DATA[idx], %_ELECTROCHEM_DATA[idx + 1], %_ELECTROCHEM_DATA[idx + 2], %_ELECTROCHEM_DATA[idx + 3]]
        )$

        chem_electrochem_half_equation(oxidant, reductant) := block([idx, n],
          idx: _electrochem_find_pair(oxidant, reductant),
          if is(idx = 0) then ""
          else (
            n: %_ELECTROCHEM_DATA[idx + 2],
            sconcat("\\ce{", oxidant, " + ", string(n), " e^- -> ", reductant, "}")
          )
        )$

        chem_electrochem_core_pair_data(oxidant, reductant) := block([ox, red],
          oxidant: _electrochem_normalize_species(oxidant),
          reductant: _electrochem_normalize_species(reductant),
          ox: ssubst("", " + 14H+", oxidant), ox: ssubst("", " + 8H+", ox),
          ox: ssubst("", " + 4H+", ox), ox: ssubst("", " + 2H+", ox),
          ox: ssubst("", "H+ + ", ox), ox: ssubst("", " + H+", ox),
          ox: ssubst("", " + 7H2O", ox), ox: ssubst("", " + 4H2O", ox),
          ox: ssubst("", " + 2H2O", ox), ox: ssubst("", "2H2O + ", ox),
          red: ssubst("", " + 14H+", reductant), red: ssubst("", " + 8H+", red),
          red: ssubst("", " + 4H+", red), red: ssubst("", " + 2H+", red),
          red: ssubst("", "H+ + ", red), red: ssubst("", " + H+", red),
          red: ssubst("", " + 7H2O", red), red: ssubst("", " + 4H2O", red),
          red: ssubst("", " + 2H2O", red), red: ssubst("", "2H2O + ", red),
          [ox, red]
        )$

        chem_electrochem_half_equation_excl_H_H2O(oxidant, reductant) := block([idx, n, core],
          idx: _electrochem_find_pair(oxidant, reductant),
          if is(idx = 0) then ""
          else (
            n: %_ELECTROCHEM_DATA[idx + 2], core: chem_electrochem_core_pair_data(oxidant, reductant),
            sconcat("\\ce{", first(core), " + ", string(n), " e^- -> ", second(core), "}")
          )
        )$

        chem_electrochem_Hplus_species(oxidant, reductant) := block([pair],
          oxidant: _electrochem_normalize_species(oxidant),
          pair: chem_electrochem_core_pair_data(oxidant, reductant),
          if _electrochem_contains(oxidant, "H+") then first(pair) else ""
        )$

        chem_electrochem_Hplus_count(oxidant, reductant) := block([x],
          x: ssubst("", " ", _electrochem_normalize_species(oxidant)),
          if _electrochem_contains(x, "14H+") then 14
          else if _electrochem_contains(x, "8H+") then 8
          else if _electrochem_contains(x, "4H+") then 4
          else if _electrochem_contains(x, "2H+") then 2
          else if _electrochem_contains(x, "H+") then 1
          else 0
        )$

        chem_electrochem_oxidant_array() := block([result, i, n],
          result: [], n: length(%_ELECTROCHEM_DATA), i: 1,
          while i <= n do (result: cons(%_ELECTROCHEM_DATA[i], result), i: i + 4),
          reverse(result)
        )$

        chem_electrochem_reductant_array() := block([result, i, n],
          result: [], n: length(%_ELECTROCHEM_DATA), i: 1,
          while i <= n do (result: cons(%_ELECTROCHEM_DATA[i + 1], result), i: i + 4),
          reverse(result)
        )$

        chem_electrochem_pairs_array() := block([result, i, n],
          result: [], n: length(%_ELECTROCHEM_DATA), i: 1,
          while i <= n do (result: cons([%_ELECTROCHEM_DATA[i], %_ELECTROCHEM_DATA[i + 1]], result), i: i + 4),
          reverse(result)
        )$

        chem_electrochem_strong_oxidant_array(E_min) := block([result, i, n, E0],
          result: [], n: length(%_ELECTROCHEM_DATA), i: 1,
          while i <= n do (
            E0: %_ELECTROCHEM_DATA[i + 3],
            if numberp(E0) and is(E0 >= E_min) then result: cons(%_ELECTROCHEM_DATA[i], result),
            i: i + 4
          ),
          reverse(result)
        )$

        chem_electrochem_strong_reductant_array(E_max) := block([result, i, n, E0],
          result: [], n: length(%_ELECTROCHEM_DATA), i: 1,
          while i <= n do (
            E0: %_ELECTROCHEM_DATA[i + 3],
            if numberp(E0) and is(E0 <= E_max) then result: cons(%_ELECTROCHEM_DATA[i + 1], result),
            i: i + 4
          ),
          reverse(result)
        )$

        chem_electrochem_cell_potential(cathode_ox, cathode_red, anode_ox, anode_red) := block([E_cath, E_an],
          E_cath: chem_electrochem_E0(cathode_ox, cathode_red),
          E_an: chem_electrochem_E0(anode_ox, anode_red),
          if is(E_cath = null) or is(E_an = null) then null else float(E_cath - E_an)
        )$

        chem_electrochem_spontaneous(cathode_ox, cathode_red, anode_ox, anode_red) := block([Ecell],
          Ecell: chem_electrochem_cell_potential(cathode_ox, cathode_red, anode_ox, anode_red),
          if is(Ecell = null) then false else is(Ecell > 0)
        )$

        chem_electrochem_array() := chem_electrochem_oxidant_array()$

stackfltfmt:"~f";

bool_to_score(b) := if is(b) then 1 else 0;
not_null_to_score(x) := if is(x = null) then 0 else 1;
not_empty_to_score(x) := if sequal(sconcat(x), "") then 0 else 1;
passfail(s) := if is(s = 1) then "PASS" else "FAIL";

ox_fe:"Fe^{3+}";
red_fe:"Fe^{2+}";
ox_zn:"Zn^{2+}";
red_zn:"Zn";

/* Basic pair extraction */
t_1: chem_electrochem_E0(ox_fe, red_fe);
t_2: chem_electrochem_n(ox_fe, red_fe);
t_3: chem_electrochem_E0(ox_zn, red_zn);
t_4: chem_electrochem_reductant(ox_fe);
t_5: chem_electrochem_oxidant(red_fe);
t_6: chem_electrochem_pair_data(ox_fe, red_fe);
t_7: chem_electrochem_half_equation(ox_fe, red_fe);
t_17: chem_electrochem_half_equation_excl_H_H2O("MnO4^- + 8H^+", "Mn^{2+} + 4H2O");
t_18: chem_electrochem_core_pair_data("MnO4^- + 8H^+", "Mn^{2+} + 4H2O");
t_19: chem_electrochem_Hplus_species("MnO4^- + 8H^+", "Mn^{2+} + 4H2O");
t_20: chem_electrochem_Hplus_count("MnO4^- + 8H^+", "Mn^{2+} + 4H2O");
t_21: chem_electrochem_Hplus_count(ox_fe, red_fe);

/* Navigation */
t_8: chem_electrochem_oxidant_array();
t_9: chem_electrochem_reductant_array();
t_10: chem_electrochem_pairs_array();
t_11: chem_electrochem_strong_oxidant_array(1.0);
t_12: chem_electrochem_strong_reductant_array(-2.0);

/* Cell calculations: Zn|Zn2+ || Fe3+|Fe2+ */
t_13: chem_electrochem_cell_potential(ox_fe, red_fe, ox_zn, red_zn);
t_14: chem_electrochem_spontaneous(ox_fe, red_fe, ox_zn, red_zn);

/* Edge and alias */
t_15: chem_electrochem_E0("X", "Y");
t_16: chem_electrochem_array();

s_1: bool_to_score(t_1 = 0.77);
s_2: bool_to_score(t_2 = 1);
s_3: bool_to_score(t_3 = -0.76);
s_4: bool_to_score(sequal(t_4, red_fe));
s_5: bool_to_score(sequal(t_5, ox_fe));
s_6: bool_to_score(length(t_6) = 4);
s_7: not_empty_to_score(t_7);
s_8: bool_to_score(length(t_8) > 0);
s_9: bool_to_score(length(t_9) > 0);
s_10: bool_to_score(length(t_10) > 0);
s_11: bool_to_score(member("F2", t_11));
s_12: bool_to_score(member("Na", t_12));
s_13: bool_to_score(float(round(100*t_13))/100 = 1.53);
s_14: bool_to_score(t_14 = true);
s_15: bool_to_score(t_15 = null);
s_16: bool_to_score(length(t_16) = length(t_8));
s_17: bool_to_score(sequal(t_17, "\\ce{MnO4^- + 5 e^- -> Mn^{2+}}"));
s_18: bool_to_score(length(t_18) = 2 and sequal(first(t_18), "MnO4^-") and sequal(second(t_18), "Mn^{2+}"));
s_19: bool_to_score(sequal(t_19, "MnO4^-"));
s_20: bool_to_score(t_20 = 8);
s_21: bool_to_score(t_21 = 0);

max_score:21;
total_score:s_1+s_2+s_3+s_4+s_5+s_6+s_7+s_8+s_9+s_10+s_11+s_12+s_13+s_14+s_15+s_16+s_17+s_18+s_19+s_20+s_21;
score_pct:round(float(100*total_score/max_score));

p_1:passfail(s_1); p_2:passfail(s_2); p_3:passfail(s_3); p_4:passfail(s_4);
p_5:passfail(s_5); p_6:passfail(s_6); p_7:passfail(s_7); p_8:passfail(s_8);
p_9:passfail(s_9); p_10:passfail(s_10); p_11:passfail(s_11); p_12:passfail(s_12);
p_13:passfail(s_13); p_14:passfail(s_14); p_15:passfail(s_15); p_16:passfail(s_16);
p_17:passfail(s_17); p_18:passfail(s_18);
p_19:passfail(s_19); p_20:passfail(s_20); p_21:passfail(s_21);

ta1:10;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@score_pct@}% of electrochemistry module tests pass.</text>
    </questionnote>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Wrong answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>2</showvalidation>
      <options>hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Confirm the test question</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty>0</truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Module test completed.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty>0.1</falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Enter {@ta1@} to confirm this diagnostic question.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Teacher answer</description>
      <testinput>
        <name>ans1</name>
        <value>10</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
